# Python基础教程演示

这是一个基于nbgrader的Python教学演示。

## 学习目标
- 理解变量的使用
- 掌握基本的数据类型
- 学会使用print函数

## 练习1: 变量定义

请定义一个变量`name`，并赋值为你的姓名：

<!-- grade_id: exercise_1, points: 0, grade: false, solution: true -->
```python
# 请在这里定义变量name
# YOUR CODE HERE
name = "请输入你的姓名"
```

<!-- grade_id: test_exercise_1, points: 5, grade: true, solution: false -->
```python
# 测试用例
assert isinstance(name, str), "name应该是字符串类型"
assert len(name) > 0, "name不能为空"
print(f"✓ 测试通过! 你好，{name}！")
```